Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Sreya Dhar"
COLLABORATORS = ""
```

---

<!-- grade_id: cell-8abe173ecc588e38, points: 0, grade: false, solution: true -->
```python
# Ex1
def average_word_len_in_sentence(sent):
    # Complete this function to calculate the average
    # word length in a sentence
    # Input: sentence
    # Output: average word length in sentence
    # Hint: count punctuations with whatever word they are `touching`
    # Hint: round the average to two decimal places
   
    # YOUR CODE HERE
    sentences = sent.split('.,')
    sentences = [sent.split() for sentence in sentences if len(sentence)]
    average = [sum(len(word) for word in sent)/len(sent) for sent in sentences]
    ave = average[0]
    return round(ave, 2)

average_word_len_in_sentence('This is a test sentence!')
```

Output:
```
5.46
```

<!-- grade_id: cell-88eb75a67af14058, points: 2, grade: true, solution: false -->
```python
assert average_word_len_in_sentence('This is a test sentence!') == 4
```

<!-- grade_id: cell-9a48ceb083442619, points: 2, grade: true, solution: false -->
```python
# Hidden test
```

<!-- grade_id: cell-9cd6ed6844d07523, points: 0, grade: false, solution: true -->
```python
#Ex2
def wc(filename):
    # Complete this function to count the number of lines, words, and chars in a file. 
    # Input: filename
    # Output: a tuple with line count, word count, and char count -- in this order

    # YOUR CODE HERE
    num_lines = 0
    num_words = 0
    num_chars = 0

    with open(filename, 'r') as f:
        data = f.read()
        line = data.splitlines()
        words = data.split()
        spaces = data.split(" ")
        charc = len(data) 
        
        
    return len(line), len(words), charc 
        
        
wc('ex2_data.txt')
```

Output:
```
(50, 643, 4298)
```

<!-- grade_id: cell-9784281c8327eef5, points: 4, grade: true, solution: false -->
```python
assert wc('ex2_data.txt') == (50, 643, 4298)

```

<!-- grade_id: cell-4cbfe85e128613b1, points: 0, grade: false, solution: true -->
```python
#Ex3a
def is_leap(year):
    # Complete this function to check if year is a leap year
    # Input: year
    # Output: True or False (Boolean)

    # YOUR CODE HERE
        # YOUR CODE HERE
        leapy = False
    
    # logic here
        if year % 400 == 0:
            leapy = True
        elif year % 100 == 0:
            leapy = False
        elif year % 4 == 0:
            leapy = True
    
        return leapy

is_leap(2000)
```

Output:
```
True
```

<!-- grade_id: cell-b9a05e2e268a2c6a, points: 4, grade: true, solution: false -->
```python
assert is_leap(1800) == False
assert is_leap(1900) == False
assert is_leap(1600) == True
assert is_leap(2000) == True
```

<!-- grade_id: cell-8b0875e64766bf5f, points: 0, grade: false, solution: true -->
```python
#Ex3b
import datetime 
def is_date_valid(month, day, year):
    # Complete this function to check if a data is valid, given month, day, and year.  
    # For example, 5/24/1962 is valid, but 9/31/2000 is not
    # Inputs: month, day, year
    # Output: True or False (Boolean)
    # Hint: Use is leap year function

    # YOUR CODE HERE
    isValidDate = True
      
    try : 
        datetime.datetime(int(year),  
                          int(month), int(day)) 
          
    except ValueError : 
        isValidDate = False
    
    return isValidDate
          

is_date_valid(5, 24, 1962)
```

Output:
```
True
```

<!-- grade_id: cell-0af99a4a4a8d91ee, points: 2, grade: true, solution: false -->
```python
assert is_date_valid(5, 24, 1962) == True
assert is_date_valid(9, 31, 2000) == False
```

<!-- grade_id: cell-72635d6da9df7e2b, points: 2, grade: true, solution: false -->
```python
# Hidden tests
```

<!-- grade_id: cell-cdb696c1bc3474fa, points: 0, grade: false, solution: true -->
```python
#Ex4
def is_leap(month, day, year):
        leap = False
        if year % 4 ==0:
            if year % 100 !=0 or year % 400 == 0:
                leap = True
        return leap

def is_date_valid(month, day, year):
        is_valid = True
        thirtyone_day_months = [1,3,5,7,8,10,12]
        #print(month, day, year)

        month =int(month)
        day = int(day)
        year = int(year)

        if month >12 or month <1 or day >31 or day <1 or year <1:
            is_valid = False
        elif month not in thirtyone_day_months and day ==31:
            is_valid = False

        elif is_leap(month, day, year) and month ==2 and day ==30:
            is_valid = False
        elif not is_leap(month, day, year) and month ==2 and day > 28:
            is_valid = False
        else:
            is_valid = True
            
        return  is_valid    

def day_number(month, day, year):
    # Complete this function to calculate the day_number given month, day, and year.
    # Information: The days of the year are often numbered from 1 through 365 (or 366).
    # This number can be computed in three steps using int arithmetic:
    # (a) - day_num = 31 * (month - 1) + day
    # (b) - if the month is after February subtract (4*(month)+23)//10
    # (c) - if it's a leap year and after February 29, add 1
    # Hint: First verify the input date is valid, return False if it is not valid; use is_date_valid
    # Hint: Also use the is_leap function
    # Inputs: month, day, year
    # Output: the day number or False (boolean) if the date is invalid. 
    
    # YOUR CODE HERE
    if not is_date_valid(month, day, year):
        return False
    
    elif (month > 2 and year % 4 == 0 and 
       (year % 100 != 0 or year % 400 == 0)): 
        day += 1; 
    
    days = [31, 28, 31, 30, 31, 30, 
        31, 31, 30, 31, 30, 31];
    # Add the days in the previous months 
    month -= 1; 
    while (month > 0): 
        day = day + days[month - 1]; 
        month -= 1; 
    return day; 
    

day_number(2, 29, 1800)
```

Output:
```
False
```

<!-- grade_id: cell-411ea3b17d1adb33, points: 2, grade: true, solution: false -->
```python
assert day_number(9, 31, 2000) == False
assert day_number(2, 13, 2020) == 44
```

<!-- grade_id: cell-e03020ec7bc1b440, points: 0, grade: false, solution: true -->
```python
#Ex5
def years_to_double_investment(apr):
    # Complete this function to use a while loop to determine how long it takes for an investment 
    # to double at a given interest rate. The input to this function, apr, is the annualized interest rate
    # and the output is the number of years it takes an investment to double. Note: The amount of the initial 
    # investment (principal) does not matter; you can use $1. 
    # Hint: principal is the amount of money being invested. 
    # apr is the annual percentage rate expressed as a decimal number.  
    # Relationship: value after one year is given by principal * (1+ apr)
    principle = 1
    #rate = input("Please enter the annual interest rate: ")
    year = 0
    total = principle
    
    while total <= principle * 2 :
        
        total = total + (total * float(apr))
        year = year + 1

    return year


years_to_double_investment(0.1)
```

Output:
```
8
```

<!-- grade_id: cell-37833043ebab338a, points: 2, grade: true, solution: false -->
```python
assert years_to_double_investment(0.06) == 12
assert years_to_double_investment(0.1) == 8
```

<!-- grade_id: cell-00925de994310b09, points: 2, grade: true, solution: false -->
```python
# Hidden tests

```

```python
years_to_double_investment(0.1)
```

<!-- grade_id: cell-915174ac3c531b23, points: 0, grade: false, solution: true -->
```python
#Ex6 
def stopping_time(n):
    # Complete this function to return the number of steps taken to reach 1 in
    # the Collatz sequence (https://en.wikipedia.org/wiki/Collatz_conjecture)
    step_seq = [n]
    if n < 1:
        return []
    while n > 1:
        if n % 2 == 0:
             n = n / 2
        else:
             n = 3 * n + 1
    # Added line
        step_seq.append(n)    
    return len(step_seq)-1

stopping_time(27)
```

Output:
```
111
```

<!-- grade_id: cell-cb2141fe7de88db3, points: 2, grade: true, solution: false -->
```python
assert stopping_time(12) == 9
assert stopping_time(27) == 111
```

<!-- grade_id: cell-a6ca4ec347d83750, points: 2, grade: true, solution: false -->
```python
# hidden tests
```

<!-- grade_id: cell-d5c69bc0cb4c6837, points: 0, grade: false, solution: true -->
```python
#Ex7
def is_prime(n):
    # A positive whole number > 2 is prime if no number between 2 and sqrt(n)
    # (include) evenly divides n. Write a program that accepts a value of n as
    # input and determine if the value is prime. If n is not prime, your program should
    # return False (boolean) as soon as it finds a value that evenly divides n.
    # Hint: if number is 2, return False

    import math
    
    # YOUR CODE HERE
    if n <= 2:
        return False

    i = 2
    for i in range(2, int(math.sqrt(n) + 1)):
        if (n % i == 0):
            return False

    return True

is_prime(25)
```

Output:
```
False
```

<!-- grade_id: cell-da18c35e0a5d579e, points: 4, grade: true, solution: false -->
```python
assert is_prime(2) == False
assert is_prime(3) == True
assert is_prime(5) == True
assert is_prime(25) == False

```

<!-- grade_id: cell-ad04b1135365e292, points: 0, grade: false, solution: true -->
```python
#Ex8
def all_primes(n):
#     # Complete this function to return all the primes as a list less than or equal to n
#     # Input: n
#     # Output: a list of numbers

#     # YOUR CODE HERE

    out = list()
    sieve = [True] * (n+1)
    for m in range(2, n+1):
        if (sieve[m]):
            out.append(m)
            for i in range(m, n+1, m):
                sieve[i] = False
    x=len(out)
    return out[-x+1:]

all_primes(5)
```

Output:
```
[3, 5]
```

<!-- grade_id: cell-3e7dbc25ba03421c, points: 2, grade: true, solution: false -->
```python
assert all_primes(5) == [3, 5]
assert all_primes(25) == [3, 5, 7, 11, 13, 17, 19, 23]

```

<!-- grade_id: cell-d6cdbdb6ba20e916, points: 0, grade: false, solution: true -->
```python
#Ex9
def gcd(m,n):
    # Complete this function to determine the greatest common divisor (GCD). 
    # The GCD of two values can be computed using Euclid's algorithm. Starting with the values
    # m and n, we repeatedly apply the formula: n, m = m, n%m until m is 0. At this point, n is the GCD
    # of the original m and n. 
    # Inputs: m and n which are both natural numbers
    # Output: gcd

    # YOUR CODE HERE
    while n:
        m,n = n, m%n
    return m

gcd(25,75)
gcd(3,13)
```

Output:
```
1
```

<!-- grade_id: cell-9a4f50a09e28d52d, points: 2, grade: true, solution: false -->
```python
assert gcd(25,75) == 25
assert gcd(3,13) == 1
```

<!-- grade_id: cell-6bba97d2c8fa2d9f, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-2e1d0f8cbd0fefe2, points: 0, grade: false, solution: true -->
```python
#Ex10
import statistics
def determine_min_max_average(filename):
    # Complete this function to read grades from a file and determine the student with the highest average 
    # test grades and the lowest average test grades. 
    # Input: filename
    # Output: a tuple containing four elements: name of student with highest average, their average, 
    # name of the student with the lowest test grade, and their average. Example ('Student1', 99.50, 'Student5', 65.50)
    # Hint: Round to two decimal places

    # YOUR CODE HERE
    
    file =open(filename)
    ave_list = []
    name_list = []
    for line in file:
        if line.strip('\n'):
            line = line.strip('\n').split(',')
            name = line[0]
            #print(name)
            grades = [int(i) for i in line[1:] ]
            #print(grades)
            average = statistics.mean(grades)
            #print(average)
            ave_list.append(average)
            name_list.append(name)
    #print(name_list)
    average_max = max(ave_list)
    average_min = min(ave_list)
    indi_max = ave_list.index(average_max)
    indi_min = ave_list.index(average_min)
    name_max = name_list[int(indi_max)]
    name_min = name_list[int(indi_min)]
    
    
    
            
        #print(average)
    return name_max, average_max, name_min, average_min 
        


determine_min_max_average('ex10_data.txt')
```

Output:
```
('student733', 86.2, 'student202', 65.4)
```

<!-- grade_id: cell-c8ec04f59adc84af, points: 4, grade: true, solution: false -->
```python
assert determine_min_max_average('ex10_data.txt') == ('student733', 86.2, 'student202', 65.4)
```

<!-- grade_id: cell-423f49b4dbb9cda6, points: 0, grade: false, solution: true -->
```python
def is_license_plate_valid(plate):

    # In a particular jurisdiction, older license plates consist of three uppercase
    # letters followed by three digits. When all of the license plates following
    # that pattern had been used, the format was changed to four digits followed by
    # three uppercase letters. 

    # Complete this function whose only input is a license plate and its output
    # is: 1) Older/Valid 2) Newer/Valid 3) Invalid
    # input: plate (str)
    # output: 'Older/Valid' or 'Newer/Valid' or 'Invalid'

    # YOUR CODE HERE
    letters = ["A","B","C","D","E","F","G","H","I","J","K","L","M","N","O","P","Q","R","S","T","U","V","W","X","Y","Z"]
    numbers = ["1","2","3","4","5","6","7","8","9","0"]

    
    if len(plate)==7:

        if plate[0] in numbers and plate[1] in numbers and plate[2] in numbers and plate[3] in numbers and plate[4] in letters and plate[5] in letters and plate[6] in letters:
            
            return "Newer/Valid"
        
        else:
            plate not in plate.upper()
            return "Invalid"
        

    elif len(plate)==6:

        if plate[3] in numbers and plate[4] in numbers and plate[5] in numbers and plate[0] in letters and plate[1] in letters and plate[2] in letters:
            
            return "Older/Valid"
          
        else:
            plate not in plate.upper()
            return "Invalid"     
    else:
        return "Invalid"

is_license_plate_valid('bTR342')

```

Output:
```
'Invalid'
```

<!-- grade_id: cell-1900e84b3f513c64, points: 4, grade: true, solution: false -->
```python
assert is_license_plate_valid('ABC123') == 'Older/Valid'
assert is_license_plate_valid('GHE952') == 'Older/Valid'
assert is_license_plate_valid('1934ABT') == 'Newer/Valid'
assert is_license_plate_valid('bTR342') == 'Invalid'
```

<!-- grade_id: cell-cfcb59d207e5d3cd, points: 4, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-859f437c63175722, points: 0, grade: false, solution: true -->
```python

def year_last_two_digits(number, digits_count=2):
    return abs(number) % (10**digits_count)

# Ex12: Magic dates
def is_magic_date(date):
    # A magic date is a date where the day multiplied by the month is equal 
    # to the two digit year. For example, June 10, 1960 is a magic date because
    # June is the sixth month, and 6 times 10 is 60, which is equal to the two
    # digit year. Complete this function to determine whether or not a date is
    # a magic date.

    # input: date (str -- month/day/year)
    # output: True or False (bool)

    date_1= date.split("/") #('month/day/year')
        
    month = int (date_1[0])
    day = int (date_1[1])
    year = int (date_1[2])
    
    
    # print (date_1)
    
    # YOUR CODE HERE
    if day > 31 or day < 1 or month > 12 or month < 1 or (year % 100) > 99 or (year % 100) < 0:
        return False
   
    else:
        dayTimesMonth = day * month
    
        if dayTimesMonth == (year % 100):
            return True
        else:
            return False
    
is_magic_date('6/8/2020')  
```

Output:
```
False
```

<!-- grade_id: cell-7a1fa98499aba445, points: 2, grade: true, solution: false -->
```python
assert is_magic_date('6/10/1960') == True
assert is_magic_date('6/8/1948') == True
```

<!-- grade_id: cell-22929b608ed4e3ba, points: 0, grade: false, solution: true -->
```python
# Ex13
import pandas as pd
from scipy import stats
def remove_outliers(data, num_outliers):
    # When analyzing data collected as a part of a science experiment it 
    # may be desriable to remove the most extreme values before performing
    # other calculations. Complete this function which takes a list of
    # values and an non-negative integer, num_outliers, as its parameters.
    # The function should create a new copy of the list with the num_outliers 
    # largest elements and the num_outliers smallest elements removed. 
    # Then it should return teh new copy of the list as the function's only 
    # result. The order of the elements in the returned list does not have to
    # match the order of the lemetns in the original list.
    # input1: data (list)
    # input2: num_outliers (int)

    # output: list
        
    # YOUR CODE HERE
    
    for i in range(num_outliers):
        data.remove(min(data))
        data.remove(max(data))
    
    # print(data)
    return sorted(data)
    
    
    

```

<!-- grade_id: cell-f37107e0e303a5ca, points: 2, grade: true, solution: false -->
```python
import random
random.seed(1234)
data = [random.randint(50, 150) for ele in range(100)]
data[45] = 1
data[46] = 2
data[90] = 250
data[34] = 300

result = [50, 50, 51, 52, 52, 52, 53, 53, 54, 
55, 55, 55, 55, 56, 58, 58, 59, 59, 59, 60, 61, 61, 
61, 62, 64, 64, 64, 68, 68, 68, 69, 69, 69, 70, 71, 
72, 73, 75, 77, 80, 81, 81, 84, 84, 84, 85, 88, 88, 89, 
92, 94, 94, 95, 95, 98, 103, 106, 108, 108, 109, 109, 109, 
110, 111, 111, 112, 113, 114, 115, 115, 117, 117, 119, 121, 
124, 124, 125, 126, 128, 129, 132, 132, 133, 134, 135, 135, 
135, 136, 138, 140, 141, 148, 148, 149, 149, 150]

assert remove_outliers(data, 2) == result

```

<!-- grade_id: cell-f584e176be7f53bc, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-76c8326115a8aee3, points: 0, grade: false, solution: true -->
```python
# Ex14: Removing duplicates

def remove_duplicates(words):
    # Complete this function to remove duplicates from the words list using a loop
    # input: words (list)
    # output: a list without duplicates

    # YOUR CODE HERE
    res = [i for n, i in enumerate(words) if i not in words[:n]] 
    return res
            

remove_duplicates([1,2,3,3,3,4,5,6,7,7,8])
```

Output:
```
[1, 2, 3, 4, 5, 6, 7, 8]
```

<!-- grade_id: cell-0b3ad78a5f56ef26, points: 2, grade: true, solution: false -->
```python
assert remove_duplicates([1,2,3,3,3,4,5,6,7,7,8]) == [1, 2, 3, 4, 5, 6, 7, 8]

```

<!-- grade_id: cell-6f1773c6eb139b8d, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-3b7cdbf9f92b6fb9, points: 0, grade: false, solution: true -->
```python
# Ex15: List of proper divisors

def proper_divisors(n):
    # A proper divisor ofa  positive integer, n, is a positive integer less than n which divides
    # evenly into n. Complete this function to compute all the proper divisors of a positive
    # integer. The integer is passed to this function as the only parameter. The function will
    # return a list of containing all of the proper divisors as its only reult. 

    # input: n (int)
    # output: list

    # YOUR CODE HERE
    for i in range(n):
        x = [i for i in range(1,n) if not n % i]
    return x

proper_divisors(28)
```

Output:
```
[1, 2, 4, 7, 14]
```

<!-- grade_id: cell-13e6b6403196da76, points: 2, grade: true, solution: false -->
```python
assert proper_divisors(28) == [1, 2, 4, 7, 14]

```

<!-- grade_id: cell-e7223c2aeb9c5d36, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-a94951aac16e950f, points: 0, grade: false, solution: true -->
```python
#Ex16: Perfect Numbers
def is_number_perfect(n):
    # An integer, n, is said to be perfect when the sum of all of the proper divisors 
    # of n is equal to n. For example, 28 is a perfect number because its proper divisors
    # are 1, 2, 4, 7, and 14 = 28 
    # Complete this function to determine if a the number a perfect number or not. 
    # input: n (int)
    # output: True or False (bool)


    # YOUR CODE HERE
    sum = 0
    for x in range(1, n):
        if n % x == 0:
            sum += x
    return sum == n

is_number_perfect(28)
```

Output:
```
True
```

<!-- grade_id: cell-14d8cfff6ced5461, points: 2, grade: true, solution: false -->
```python
assert is_number_perfect(28) == True
assert is_number_perfect(76) == False

```

<!-- grade_id: cell-14eb44447e299251, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-21d1d1b24e27c644, points: 0, grade: false, solution: true -->
```python
# Ex17: Line of Best Fit
import numpy as np
import pandas as pd

def best_line(points):
    # Complete this function to determine the best line. 
    # https://www.varsitytutors.com/hotmath/hotmath_help/topics/line-of-best-fit
    # input: points (list of tuples contain x, y values)
    # output: (m, b) # round both values to two decimal places

    # YOUR CODE HERE
    df = pd.DataFrame(points, columns =['X', 'Y'])  
    X= df["X"].values
    Y= df["Y"].values
    # Mean X and Y
    mean_x = np.mean(X)
    mean_y = np.mean(Y)
 
    # Total number of values
    n = len(X)
    # Using the formula to calculate 'm' and 'b'
    numer = 0
    denom = 0
    for i in range(n):
        numer += (X[i] - mean_x) * (Y[i] - mean_y)
        denom += (X[i] - mean_x) ** 2
        m = numer / denom
        b = mean_y - (m * mean_x)
        
    return round(m,2), round(b,2)
    

```

<!-- grade_id: cell-386bd6a9301a19da, points: 2, grade: true, solution: false -->
```python
points = (
    (8, 3),
    (2, 10),
    (11, 3),
    (6, 6),
    (5, 8),
    (4, 12),
    (12, 1),
    (9, 4),
    (6, 9),
    (1, 14),
)

assert best_line(points) == (-1.11, 14.08)
```

```python

```

<!-- grade_id: cell-e3b40327cdd21344, points: 2, grade: true, solution: false -->
```python

```